---
title       : Insert the chapter title here
description : Insert the chapter description here


--- type:NormalExercise lang:python xp:100 skills:2 key:c6434d6e6b
## <<<New Exercise>>>


*** =instructions

*** =hint

*** =pre_exercise_code
```{python}

```

*** =sample_code
```{python}

```

*** =solution
```{python}

```

*** =sct
```{python}

```
